rule_files:
  - slo-availability.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: istio_requests_total{job="ingress-metrics-proxy",response_code="500"}
        values: "0+0x30 0+10x30"
      - series: istio_requests_total{job="ingress-metrics-proxy"}
        values: "0+1x30 31+90x30"
    alert_rule_test:
      - eval_time: 30m
        alertname: KuadrantAvailabilityHighErrorRate
        exp_alerts: []
      - eval_time: 60m
        alertname: KuadrantAvailabilityHighErrorRate
        exp_alerts:
            - exp_labels:
                  alertname: KuadrantAvailabilityHighErrorRate
                  category: availability
                  owner: kuadrant-org
                  severity: critical
                  sloth_id: kuadrant-requests-availability
                  sloth_service: kuadrant
                  sloth_severity: page
                  sloth_slo: requests-availability
              exp_annotations:
                  summary: High error rate on HTTPRoute requests responses
                  title: (page) kuadrant requests-availability SLO error budget burn rate is too fast.
            - exp_labels:
                  alertname: KuadrantAvailabilityHighErrorRate
                  category: availability
                  owner: kuadrant-org
                  severity: warning
                  sloth_id: kuadrant-requests-availability
                  sloth_service: kuadrant
                  sloth_severity: ticket
                  sloth_slo: requests-availability
              exp_annotations:
                  summary: High error rate on HTTPRoute requests responses
                  title: (ticket) kuadrant requests-availability SLO error budget burn rate is too fast.          
        