rule_files:
  - /prometheus/latency-rules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: istio_request_duration_milliseconds_bucket{le="250",job="ingress-metrics-proxy",response_code="200"}
        values: "0+1x30 31+10x30"
      - series: istio_request_duration_milliseconds_count{job="ingress-metrics-proxy",response_code="200"}
        values: "0+1x30 31+100x30"
    alert_rule_test:
      - eval_time: 30m
        alertname: KuadrantlatencyHighErrorRate
        exp_alerts: []
      - eval_time: 60m
        alertname: KuadrantlatencyHighErrorRate
        exp_alerts:
            - exp_labels:
                  alertname: KuadrantlatencyHighErrorRate
                  category: latency
                  owner: kuadrant-org
                  severity: critical
                  sloth_id: kuadrant-requests-latency
                  sloth_service: kuadrant
                  sloth_severity: page
                  sloth_slo: requests-latency
              exp_annotations:
                  summary: High latency on HTTPRoute requests responses 
                  title: (page) kuadrant requests-latency SLO error budget burn rate is too fast.
            - exp_labels:
                  alertname: KuadrantlatencyHighErrorRate
                  category: latency
                  owner: kuadrant-org
                  severity: warning
                  sloth_id: kuadrant-requests-latency
                  sloth_service: kuadrant
                  sloth_severity: ticket
                  sloth_slo: requests-latency
              exp_annotations:
                  summary: High latency on HTTPRoute requests responses 
                  title: (ticket) kuadrant requests-latency SLO error budget burn rate is too fast.       
    